<a href="https://colab.research.google.com/github/njokimainavet/GWAC-PRACTICE/blob/main/DSA_2026_Entry.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
#!pip install otter-grader
import otter
grader = otter.Notebook("DSA_2026_Entry.ipynb")

```

Output:
```
ModuleNotFoundError: No module named 'otter'

---------------------------------------------------------------------------
NOTE: If your import is failing due to a missing package, you can
manually install dependencies using either !pip or !apt.

To view examples of installing some common dependencies, click the
"Open Examples" button below.
---------------------------------------------------------------------------
```

## DSA 2026 Summer School Admittance Check

Thanks for your interest in attending DSA Kampala 2026 at Makerere University. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School.
Good luck! See you in Kampala at Makerere University!

```python
# Run these once ... Just in case
#!pip install otter-grader
#!pip install matplotlib scipy  # For machine learning visualization questions
import IPython
from IPython import get_ipython
# Import the good stuff
import pandas as pd
import numpy as np
import math
import matplotlib.pyplot as plt
from scipy.stats import multivariate_normal
# Note: grader is already initialized in Cell 0 with the notebook path
# Don't reinitialize it here, just use the existing grader
```

**The Extra Mile!**

This section will test your data analysis skills using real-world economic data from Uganda. You'll work with Consumer Price Index (CPI) data to practice data manipulation, statistical analysis, and time series handling. Good luck!

**Question E1:** Load the Uganda Consumer Price Index dataset from the file `data/uganda-consumer-price-index-trends-2020-2023.xlsx` into a pandas dataframe. Display the first 10 rows and the shape of the dataframe.

```python
import pandas as pd
import numpy as np

# Load the Uganda CPI dataset
df = ...

# Display first 10 rows and shape
print("Shape of dataframe:", ...)
print("\nFirst 10 rows:")
print(...)
```

Output:
```
Shape of dataframe: Ellipsis

First 10 rows:
Ellipsis
```

```python
grader.check("e1")
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipython-input-4074947994.py in <cell line: 0>()
----> 1 grader.check("e1")

NameError: name 'grader' is not defined
```

**Question E2:** Reshape the dataframe from wide format to long format. The monthly columns (dec-20, jan-21, etc.) should become a single 'date' column, and their values should be in a 'value' column. The resulting dataframe should have columns: 'indicator_code', 'description', 'date', and 'value'.

```python
# Reshape dataframe from wide to long format
df_long = ...

print("Shape after reshaping:", df_long.shape)
print("\nFirst 10 rows:")
print(df_long.head(10))
```

```python
grader.check("e2")
```

**Question E3:** Convert the 'date' column to datetime format. Handle the date format appropriately (e.g., 'dec-20' should become '2020-12-01', 'jan-21' should become '2021-01-01', etc.).

```python
# Convert date column to datetime
df_long['date'] = ...

print("Date column info:")
print(df_long['date'].head(10))
print("\nDate range:", df_long['date'].min(), "to", df_long['date'].max())
```

```python
grader.check("e3")
```

**Question E4:** Filter the dataframe to include only CPI indicators (indicator_code starting with 'CPI_'). Create separate dataframes for:
- All Items CPI (CPI_16 and CPI_09)
- Core CPI (CPI_CORE_16 and CPI_CORE_09)
- Food CPI (CPI_FOOD_16 and CPI_FOOD_09)
- Energy Fuel and Utilities CPI (CPI_EFU_16 and CPI_EFU_09)

```python
# Filter for CPI indicators
df_cpi = ...

# Create separate dataframes
df_all_items = ...
df_core = ...
df_food = ...
df_efu = ...

print("All Items CPI shape:", df_all_items.shape)
print("Core CPI shape:", df_core.shape)
print("Food CPI shape:", df_food.shape)
print("EFU CPI shape:", df_efu.shape)

```

```python
grader.check("e4")
```

**Question E5:** Check for missing values in the 'value' column. If there are any missing values, replace them with the median value of that specific indicator_code. Display the count of missing values before and after replacement.

```python
# Check for missing values
missing_before = ...

print("Missing values before replacement:", missing_before)

# Replace missing values with median per indicator_code
df_long['value'] = ...

missing_after = ...
print("Missing values after replacement:", missing_after)
```

```python
grader.check("e5")
```

## Natural Language Processing and Large Language Models

**Question N1:** Write a function `count_words` that takes a text string as input and returns a dictionary where keys are unique words (lowercased) and values are their frequencies. Ignore punctuation and split on whitespace.

Example: `count_words("Hello world hello")` should return `{'hello': 2, 'world': 1}`

```python
def count_words(text):
    """
    Count word frequencies in a text string.

    Parameters:
    text (str): Input text string

    Returns:
    dict: Dictionary with words as keys and frequencies as values
    """
    ...
    return word_counts

# Test the function
test_text = "Data Science Africa 2026 is happening in Kampala at Makerere University. Makerere University is a great place for learning."
result = count_words(test_text)
print(result)
```

```python
grader.check("n1")
```

**Question N2:** Write a function `extract_entities` that takes a text string and a list of entity keywords, and returns a list of sentences that contain any of those keywords. This simulates basic named entity recognition.

Example:
- Text: "Makerere University is in Kampala. Kampala is the capital of Uganda."
- Keywords: ["Makerere", "Kampala"]
- Should return: ["Makerere University is in Kampala.", "Kampala is the capital of Uganda."]

```python
def extract_entities(text, keywords):
    """
    Extract sentences containing specified keywords.

    Parameters:
    text (str): Input text string
    keywords (list): List of keywords to search for

    Returns:
    list: List of sentences containing the keywords
    """
    ...
    return matching_sentences

# Test the function
test_text = "DSA 2026 will be held at Makerere University in Kampala. Kampala is a beautiful city. Makerere University is one of the oldest universities in Africa."
keywords = ["Makerere", "Kampala", "DSA"]
result = extract_entities(test_text, keywords)
print(result)
```

```python
grader.check("n2")
```

**Question N3:** Write a function `calculate_similarity` that takes two text strings and calculates their similarity based on the Jaccard similarity of their word sets. Jaccard similarity is the size of the intersection divided by the size of the union of two sets.

Formula: J(A, B) = |A ∩ B| / |A ∪ B|

Example:
- Text1: "data science africa"
- Text2: "data science kampala"
- Should return approximately 0.67 (2 common words / 3 unique words)

```python
def calculate_similarity(text1, text2):
    """
    Calculate Jaccard similarity between two text strings.

    Parameters:
    text1 (str): First text string
    text2 (str): Second text string

    Returns:
    float: Jaccard similarity score between 0 and 1
    """
    ...
    return similarity

# Test the function
text1 = "Data Science Africa 2026 Kampala Makerere University"
text2 = "Data Science Africa Kampala Makerere"
similarity = calculate_similarity(text1, text2)
print(f"Similarity: {similarity:.4f}")
```

```python
grader.check("n3")
```

## Machine Learning and Data Visualization

This section introduces you to data visualization and basic machine learning concepts using 2D classification data.

**Question M1:** Load the dataset from `data/data-2class.npz`. This file contains a set of 2-dimensional points `d` (shape: 1000x2), and a corresponding set of labels `l` (shape: 1000x1). Create a 2D scatterplot of the points, using red for elements with label 0, and blue for elements with label 1.

**Hint:** Use matplotlib for plotting. The dataset keys are 'd' for data points and 'l' for labels.

```python
import matplotlib.pyplot as plt
import numpy as np

# Load the dataset
data = ...

# Extract points and labels
d = ...
l = ...

# Create scatterplot
# Red for label 0, blue for label 1
...

plt.xlabel('X')
plt.ylabel('Y')
plt.title('2D Scatterplot of Data Points by Class')
plt.legend(['Class 0', 'Class 1'])
plt.show()
```

```python
grader.check("m1")
```

**Question M2:** Draw a straight line separating the two classes on the scatterplot. The line should visually separate the red points (label 0) from the blue points (label 1). You can choose the line parameters (slope and intercept) manually - the purpose is to think about how this line could be used to classify the data.

**Hint:** You can use `plt.plot()` or `plt.axline()` to draw a line. Think about where the line should be positioned to best separate the two classes.

```python
# Recreate the scatterplot
...

# Draw a separating line
# You can choose the line parameters (slope and intercept)
# Example: plt.axline((x1, y1), slope=m) or plt.plot([x1, x2], [y1, y2])
...

plt.xlabel('X')
plt.ylabel('Y')
plt.title('2D Scatterplot with Separating Line')
plt.legend(['Class 0', 'Class 1', 'Separating Line'])
plt.show()
```

```python
grader.check("m2")
```

**Question M3:** Fit two 2D Gaussian distributions to the points with label $l_i=0$ and $l_i=1$. For each class, calculate:
- The mean (centroid) of the points
- The covariance matrix

**Hint:** Use NumPy to calculate the mean and covariance. For class 0, filter points where `l == 0`, and for class 1, filter points where `l == 1`.

```python
# Filter points by label
# Reshape labels if needed: l = l.flatten() or l = l.ravel()
points_class_0 = ...
points_class_1 = ...

# Calculate mean (centroid) for each class
mean_class_0 = ...
mean_class_1 = ...

# Calculate covariance matrix for each class
cov_class_0 = ...
cov_class_1 = ...

print("Class 0:")
print(f"  Mean: {mean_class_0}")
print(f"  Covariance:\n{cov_class_0}")

print("\nClass 1:")
print(f"  Mean: {mean_class_1}")
print(f"  Covariance:\n{cov_class_1}")
```

```python
grader.check("m3")
```

**Question M4:** Create a heatmap plotting the two Gaussian distributions and superimpose a scatterplot of the data points.

**Hint:**
- Use `scipy.stats.multivariate_normal` to create the Gaussian distributions
- Create a meshgrid for the heatmap
- Use `plt.contour()` or `plt.contourf()` for the heatmap
- Overlay the scatterplot on top

```python
from scipy.stats import multivariate_normal

# Create meshgrid for heatmap
x_min, x_max = d[:, 0].min() - 1, d[:, 0].max() + 1
y_min, y_max = d[:, 1].min() - 1, d[:, 1].max() + 1
xx, yy = np.meshgrid(np.linspace(x_min, x_max, 100),
                     np.linspace(y_min, y_max, 100))

# Create Gaussian distributions
gaussian_class_0 = ...
gaussian_class_1 = ...

# Calculate probability density for each point in meshgrid
pos = np.dstack((xx, yy))
pdf_class_0 = ...
pdf_class_1 = ...

# Create heatmap (contour plot)
...

# Superimpose scatterplot
...

plt.xlabel('X')
plt.ylabel('Y')
plt.title('Gaussian Distributions Heatmap with Data Points')
plt.colorbar(label='Probability Density')
plt.show()
```

```python
grader.check("m4")
```

## Submission Instructions

### Before Submitting:

1. **Run all cells in order** - Make sure you've executed every cell from top to bottom. This ensures all your solutions are saved and any outputs are included.

2. **Check all tests pass** - Verify that all `grader.check()` cells show that tests pass. If any fail, fix your code and re-run.

3. **Save your notebook** - Use File → Save (or Cmd+S / Ctrl+S) to save all your work.

4. **Verify your solutions** - Double-check that:
   - All `...` placeholders have been replaced with your code
   - All functions are properly implemented
   - Data analysis questions produce expected outputs
   - NLP functions work with the test cases

### Submitting:

1. **Run the cell below** - This will generate a `.zip` file containing your completed notebook.

2. **Find the zip file** - It will be created in the same directory as this notebook, typically named something like `DSA_2026_Entry.zip` or similar.

3. **Upload to application portal** - Upload the generated zip file to:
   - The Google Sheet link provided in the application instructions, OR
   - The application portal/website as specified

4. **Keep a backup** - Save a copy of your completed notebook for your records.

### Important Reminders:

- ⚠️ **All work must be your own** - Uploading work that is not yours will lead to non-admittance to DSA 2026 Kampala at Makerere University.
- ✅ Make sure the zip file contains your completed notebook with all solutions
- ✅ Do not modify the test files or grader setup
- ✅ Include the `data/` folder if required by the submission instructions

**Ready to submit? Run the cell below!**

```python
# Save your notebook first, then run this cell to export your submission.grader.export(run_tests=True)
```

```python

```

Tests for q1 (points: 2)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_primes(sieve):
...     assert sieve(30) == ((2, 3, 5, 7, 11, 13, 17, 19, 23, 29), (4, 6, 8, 9, 10, 12, 14, 15, 16, 18, 20, 21, 22, 24, 25, 26, 27, 28, 30))
>>> test_primes(sieve)
```


Tests for q2 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_dictionary(Counter):
...     assert Counter(paragraph) == {' ': 192, 'e': 89, 'o': 68, 't': 65, 'i': 60, 's': 57, 'a': 48, 'r': 45, 'n': 42, 'l': 35, 'c': 31, 'd': 27, 'f': 26, 'u': 26, '0': 22, 'm': 20, 'y': 19, 'h': 17, 'w': 16, 'g': 14, '3': 13, '.': 12, 'p': 11, 'b': 11, '1': 11, '2': 10, '7': 10, 'k': 9, ',': 7, '4': 5, 'v': 4, '-': 4, '*': 4, '^': 4, 'T': 3, 'x': 3, '5': 3, '6': 3, '8': 3, '’': 3, '(': 3, ')': 3, 'q': 3, 'O': 3, ':': 2, '9': 2, 'S': 2, 'A': 2, '{': 2, '}': 2, '`': 2, 'I': 1, 'V': 1, 'F': 1, '[': 1, ']': 1, 'W': 1, '<': 1, 'N': 1, '>': 1, '“': 1, '”': 1, '‘': 1, 'E': 1, '/': 1, '\\': 1}
```


Tests for q3 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_keyslist():
...     assert list(dictionary.keys()) == ['T', 'h', 'e', ' ', 'f', 'o', 'l', 'w', 'i', 'n', 'g', 's', 'a', 'm', 'p', 't', 'x', 'I', 'u', 'd', 'r', 'c', 'k', 'y', 'b', '.', ':', '1', '2', '3', '4', '5', '6', '7', '8', '9', 'S', 'V', '’', '0', ',', '(', ')', 'q', 'A', 'v', 'F', '-', '{', '}', 'O', '[', ']', 'W', '<', 'N', '>', '`', '*', '^', '“', '”', '‘', 'E', '/', '\\']
>>> test_keyslist()
```


Tests for q4 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_valueslist():
...     assert list(dictionary.values()) == [3, 17, 89, 192, 26, 68, 35, 16, 60, 42, 14, 57, 48, 20, 11, 65, 3, 1, 26, 27, 45, 31, 9, 19, 11, 12, 2, 11, 10, 13, 5, 3, 3, 10, 3, 2, 2, 1, 3, 22, 7, 3, 3, 3, 2, 4, 1, 4, 2, 2, 3, 1, 1, 1, 1, 1, 1, 2, 4, 4, 1, 1, 1, 1, 1, 1]
```


Tests for q5 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_tuple():
...     assert final_tuple == (('T', 'h', 'e', ' ', 'f', 'o', 'l', 'w', 'i', 'n', 'g', 's', 'a', 'm', 'p', 't', 'x', 'I', 'u', 'd', 'r', 'c', 'k', 'y', 'b', '.', ':', '1', '2', '3', '4', '5', '6', '7', '8', '9', 'S', 'V', '’', '0', ',', '(', ')', 'q', 'A', 'v', 'F', '-', '{', '}', 'O', '[', ']', 'W', '<', 'N', '>', '`', '*', '^', '“', '”', '‘', 'E', '/', '\\'), (3, 17, 89, 192, 26, 68, 35, 16, 60, 42, 14, 57, 48, 20, 11, 65, 3, 1, 26, 27, 45, 31, 9, 19, 11, 12, 2, 11, 10, 13, 5, 3, 3, 10, 3, 2, 2, 1, 3, 22, 7, 3, 3, 3, 2, 4, 1, 4, 2, 2, 3, 1, 1, 1, 1, 1, 1, 2, 4, 4, 1, 1, 1, 1, 1, 1))
>>> test_tuple()
```


Tests for q6 (points: 10)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```


Tests for q8 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q9 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```
